# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

```python
import os
os.getcwd()
```

Output:
```
'C:\\Users\\airfo\\Downloads'
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

```python
Energy = pd.read_excel('Energy Indicators.xls', 
                      skiprows = 17,
                      skipfooter = 38,
                      usecols = 'C:G')
Energy.columns = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
# gets rid of missing values and replaces with NaN.
Energy = Energy.replace('...', np.NaN)
# multiples the column values by 1M
Energy["Energy Supply"] = 1000000 * Energy["Energy Supply"]
# gets rid of footnote numbers in country names
Energy['Country'] = Energy['Country'].str.replace('\d+', '')
Energy['Country'] = Energy['Country'].str.replace(r'\([^)]*\)', '')
Energy['Country'] = Energy['Country'].str.strip()
# replaces values with specified value
Energy['Country'] = Energy['Country'].replace(["Republic of Korea","United States of America",
                                               "China, Hong Kong Special Administrative Region",
                                              "United Kingdom of Great Britain and Northern Ireland",
                                              "Switzerland17","Bolivia (Plurinational State of)"],
                                              ["South Korea","United States","United Kingdom","Hong Kong","Switzerland",
                                              "Bolivia"])
# display all Rows
pd.set_option('display.max_rows', 250)
Energy.head(15)
```

Output:
```
Country  Energy Supply  Energy Supply per Capita  % Renewable
0           Afghanistan   3.210000e+08                      10.0    78.669280
1               Albania   1.020000e+08                      35.0   100.000000
2               Algeria   1.959000e+09                      51.0     0.551010
3        American Samoa            NaN                       NaN     0.641026
4               Andorra   9.000000e+06                     121.0    88.695650
5                Angola   6.420000e+08                      27.0    70.909090
6              Anguilla   2.000000e+06                     136.0     0.000000
7   Antigua and Barbuda   8.000000e+06                      84.0     0.000000
8             Argentina   3.378000e+09                      79.0    24.064520
9               Armenia   1.430000e+08                      48.0    28.236060
10                Aruba   1.200000e+07                     120.0    14.870690
11            Australia   5.386000e+09                     231.0    11.810810
12
```

```python
GDP = pd.read_csv('world_bank.csv',
                 header = 4)
GDP = GDP.replace({'Country Name': {"Korea, Rep.":"South Korea",
                                   "Iran, Islamc Rep.":"Iran",
                                   "Hong Kong SAR, China": "Hong Kong"
                  }})
GDP = GDP[['Country Name', '2006','2007','2008','2009','2010',
           '2011','2012','2013','2014','2015']]
GDP = GDP.rename(columns = {'Country Name': 'Country'})
GDP.head(15)
```

Output:
```
pandas\_libs\parsers.pyx in pandas._libs.parsers.TextReader.__cinit__()

pandas\_libs\parsers.pyx in pandas._libs.parsers.TextReader._setup_parser_source()

FileNotFoundError: [Errno 2] File world_bank.csv does not exist: 'world_bank.csv'
```

```python
ScimEn = pd.read_excel('scimagojr-3.xlsx')
ScimEn.head(15)
```

Output:
```
~\anaconda3\lib\site-packages\pandas\io\excel\_xlrd.py in load_workbook(self, filepath_or_buffer)
     34             return open_workbook(file_contents=data)
     35         else:
---> 36             return open_workbook(filepath_or_buffer)
     37 
     38     @property

~\anaconda3\lib\site-packages\xlrd\__init__.py in open_workbook(filename, logfile, verbosity, use_mmap, file_contents, encoding_override, formatting_info, on_demand, ragged_rows)
    109     else:
    110         filename = os.path.expanduser(filename)
--> 111         with open(filename, "rb") as f:
    112             peek = f.read(peeksz)
    113     if peek == b"PK\x03\x04": # a ZIP file

FileNotFoundError: [Errno 2] No such file or directory: 'scimagojr-3.xlsx'
```

```python
merged_df = pd.merge(ScimEn, Energy, how = 'inner', on = 'Country')
merged_df2 = pd.merge(merged_df, GDP, how = 'inner', on = 'Country')

merged_df2 = merged_df2.set_index('Country')

merged_df2.query('Rank < 16')

merged_df2.head(15)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-41-b3f68a162d37> in <module>
----> 1 merged_df = pd.merge(ScimEn, Energy, how = 'inner', on = 'Country')
      2 merged_df2 = pd.merge(merged_df, GDP, how = 'inner', on = 'Country')
      3 
      4 merged_df2 = merged_df2.set_index('Country')
      5 

NameError: name 'ScimEn' is not defined
```

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    merged_df3 = merged_df2.query('Rank < 16')
    merged_df3.head(15)
    raise NotImplementedError()
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

Output:
```
<ipython-input-34-e4f26ba1c257> in <module>
----> 1 assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"
      2 
      3 assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

<ipython-input-33-93809783ab90> in answer_one()
      1 def answer_one():
----> 2     merged_df3 = merged_df2.query('Rank < 16')
      3     merged_df3.head(15)
      4     raise NotImplementedError()

NameError: name 'merged_df2' is not defined
```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    merged_df2.head(200)
    len(merged_df2) - 15
    raise NotImplementedError()
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

Output:
```
<ipython-input-38-100e4222f246> in <module>
----> 1 assert type(answer_two()) == int, "Q2: You should return an int number!"

<ipython-input-37-f784e86af7c3> in answer_two()
      1 def answer_two():
----> 2     merged_df2.head(200)
      3     len(merged_df2) - 15
      4     raise NotImplementedError()

NameError: name 'merged_df2' is not defined
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    merged_df2['avg_column']= merged_df2.iloc[:, -11: ].mean(axis=1)
    merged_df2.sort_values(by=['avg_column'], inplace=True, ascending=False)
    avgGDP = merged_df2
    avgGDP.head(15)
    raise NotImplementedError()
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

Output:
```
<ipython-input-40-2a52ebf3c3cc> in <module>
----> 1 assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

<ipython-input-39-3a45aa587fed> in answer_three()
      1 def answer_three():
----> 2     merged_df2['avg_column']= merged_df2.iloc[:, -11: ].mean(axis=1)
      3     merged_df2.sort_values(by=['avg_column'], inplace=True, ascending=False)
      4     avgGDP = merged_df2
      5     avgGDP.head(15)

NameError: name 'merged_df2' is not defined
```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    UK_gdp = avgGDP

    sub_column = UK_gdp['2015'] - UK_gdp['2006']
    UK_gdp['uk_diff'] = sub_column

    UK_gdp.head(15)
    
    UK_gdp.iloc[5:6, 21:22]
    raise NotImplementedError()
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    merged_df2['Energy Supply per Capita'].mean()
    raise NotImplementedError()
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    merged_df3 = merged_df2.query('Rank < 16')
    merged_df3.sort_values(by=['% Renewable'], inplace=True, ascending=False)
    merged_df3.head()
    raise NotImplementedError()
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    column24 = merged_df2['ci_ratio']
    max_value = column24.max()
    max_value
    raise NotImplementedError()
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    merged_df2['pop_est'] = merged_df2['Energy Supply'] / merged_df2['Energy Supply per Capita']
    merged_df2.sort_values(by=['pop_est'], inplace=True, ascending=False)
    merged_df2.iloc[2:3, 0:0]
    raise NotImplementedError()
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    merged_df2['cite_pc'] =  merged_df2['Citable documents']/merged_df2['pop_est']
    raise NotImplementedError()
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

Output:
```
<ipython-input-47-e27042e076fa> in <module>
----> 1 assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

<ipython-input-45-f204783b98b3> in answer_nine()
      1 def answer_nine():
----> 2     merged_df2['cite_pc'] =  merged_df2['Citable documents']/merged_df2['pop_est']
      3     raise NotImplementedError()

NameError: name 'merged_df2' is not defined
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```